```python
version = "REPLACE_PACKAGE_VERSION"
```

---
# Assignment 1 Part 3: Hidden Markov Models (HMMs) (30 pts)

In this assignment, we're going to train a Hidden Markov Model (HMM) that is able to tag words with their [part-of-speech](https://en.wikipedia.org/wiki/Part_of_speech) (POS) in a sentence. 

```python
# Configure nltk

import nltk
from mads.lib.path import assets

nltk_data_path = assets.find("nltk_data")
if nltk_data_path not in nltk.data.path:
    nltk.data.path.append(nltk_data_path)
```

## Question 1: Load the dataset (20 pts)

Let's first load some POS-tagged data for training and testing our HMM. **Complete the `load_data` function below** that takes two arguments, `data_file`, which specifies the data file to read, and `label`, which indicates whether to include labels in your data.

Under the `conll03` subfolder in the `data` folder, you are given three data files, `eng.train`, `eng.testa` and `eng.testb`, that come from a [CoNLL-2003 shared task](https://www.clips.uantwerpen.be/conll2003/ner/). The shared task was originally held as a competition for Named Entity Recognition (NER), but the data it provided also includes POS tags that make POS Tagging possible. NER and POS Tagging are very similar tasks and HMMs are capable of handling both of them. After finishing this assignment, you are encouraged to re-purpose your HMM tagger for NER as a self-exercise. 

All three data files share the same format. They are concatenation of several documents demarcated by `-DOCSTART- -X- O O`. Sentences are separated by empty lines and each token occupies one line with its POS tag. For example, 

`scientists NNS I-NP O`

says the token `scientists` has a POS tag of `NNS` (noun, plural). If interested, see [here](https://www.ling.upenn.edu/courses/Fall_2003/ling001/penn_treebank_pos.html) for a reference list of all POS tags. The [CoNLL-2003 shared task](https://www.clips.uantwerpen.be/conll2003/ner/) page also includes some descriptions of the data. 

In many NLP tasks, it's commonplace to divide a dataset into training, development and testing sets. Since training and development sets are used for model learning and hyper-parameter tuning, they often include labels (if any). However, labels are often *not* distributed with testing sets to prevent anyone from cheating, especially in competitions. Here we will follow the same convention, by taking `eng.train` as the training set, `eng.testa` as the development set and `eng.testb` as the testing set, and **will only allow the training and development sets to include labels** by specifying the `label` argument: 

* **When called with `label=True`**, your function should return a `list`, where each element is a `list` of `tuple` representing each sentence. Each `tuple` contains a token and its POS tag, such as `("scientists", "NNS")`. 

* **When called with `label=False`**, your function should also return a `list`, where each element is just a `list` of tokens. POS tags are not included and `tuple` is not needed.

Finally, **in all the three data files, we only consider "substantial" sentences that have at least $10$ tokens**. You'd need to filter out shorter sentences and should not include them in any of your three subsets of the data. 

**This function should return a `list`, where each element is either a `list` of `tuple` or a `list` of `str`.**

**When the argument `label=True`, an example output would be:**
```
[[('The', 'DT'), ('European', 'NNP'), ('Commission', 'NNP'), ('said', 'VBD'), ('on', 'IN'), ('Thursday', 'NNP'), ('it', 'PRP'), ('disagreed', 'VBD'), ('with', 'IN'), ('German', 'JJ'), ('advice', 'NN'), ('to', 'TO'), ('consumers', 'NNS'), ('to', 'TO'), ('shun', 'VB'), ('British', 'JJ'), ('lamb', 'NN'), ('until', 'IN'), ('scientists', 'NNS'), ('determine', 'VBP'), ('whether', 'IN'), ('mad', 'JJ'), ('cow', 'NN'), ('disease', 'NN'), ('can', 'MD'), ('be', 'VB'), ('transmitted', 'VBN'), ('to', 'TO'), ('sheep', 'NN'), ('.', '.')], 
...
[('Africa', 'NNP'), (')', ')'), ('71', 'CD'), ('74', 'CD'), ('75', 'CD'), (',', ','), ('David', 'NNP'), ('Gilford', 'NNP'), ('69', 'CD'), ('74', 'CD'), ('77', 'CD'), ('.', '.')]]
```

**When the argument `label=False`, an example output would be:**

```
[['SOCCER', '-', 'JAPAN', 'GET', 'LUCKY', 'WIN', ',', 'CHINA', 'IN', 'SURPRISE', 'DEFEAT', '.'], 
 ...
 ['The', 'lanky', 'former', 'Leeds', 'United', 'defender', 'did', 'not', 'make', 'his', 'England', 'debut', 'until', 'the', 'age', 'of', '30', 'but', 'eventually', 'won', '35', 'caps', 'and', 'was', 'a', 'key', 'member', 'of', 'the', '1966', 'World', 'Cup', 'winning', 'team', 'with', 'his', 'younger', 'brother', ',', 'Bobby', '.']]
```

<!-- grade_id: cell-876ee5c9fdc4f347, points: 0, grade: false, solution: true -->
```python
from pathlib import Path

def _resolve_data_file(data_file):
    """
    Resolve a CoNLL file path from either a direct path or the course assets helper.
    """
    path = Path(str(data_file))
    if path.exists():
        return path
    
    # try common course-asset locations so the function is not tied to one exact prefix
    candidates = [
        str(data_file),
        f"conll03/{data_file}",
        f"data/conll03/{data_file}",
    ]
    
    for candidate in candidates:
        try:
            resolved = Path(assets.find(candidate))
            if resolved.exists():
                return resolved
        except Exception:
            # keep trying the other possible asset paths
            pass
    
    raise FileNotFoundError(f"could not locate data file: {data_file}")


def load_data(data_file, label=True):
    """
    Load tokens and optional POS labels from a CoNLL-2003 data file.
    """
    tagged_sents = []
    current_sent = []
    data_path = _resolve_data_file(data_file)
    
    with open(data_path, "r", encoding = "utf-8") as f:
        for line in f:
            line = line.strip()
            
            # blank lines mark sentence boundaries in the CoNLL files
            if not line:
                if len(current_sent) >= 10:
                    tagged_sents.append(current_sent)
                current_sent = []
                continue
            
            # document markers are metadata, not sentence tokens
            if line.startswith("-DOCSTART-"):
                continue
            
            parts = line.split()
            token = parts[0]
            pos_tag = parts[1]
            
            # keep POS tags for train/dev, but remove them for the held-out test set
            if label:
                current_sent.append((token, pos_tag))
            else:
                current_sent.append(token)
    
    # include the final sentence if the file does not end with a blank line
    if len(current_sent) >= 10:
        tagged_sents.append(current_sent)
    
    return tagged_sents

```

**Commentary:** The loader treats each non-empty CoNLL row as one token record and uses blank lines as sentence boundaries. It skips `-DOCSTART-` rows because they are document metadata rather than real words. The function returns `(token, POS)` pairs when labels are allowed and plain tokens when labels should be hidden. Sentences shorter than 10 tokens are filtered out to match the assignment requirement.


<!-- grade_id: cell-7896bc924dc5f94b, points: 0, grade: false, solution: true -->
```python
# Let's create the training, development and testing sets using load_data above
# train/dev keep labels, while test only keeps tokens

dataset = {
    "train": load_data("eng.train", label = True),
    "dev": load_data("eng.testa", label = True),
    "test": load_data("eng.testb", label = False),
}

```

**Commentary:** The dataset dictionary follows the standard NLP split convention: the training and development sets include labels because they are used for learning and validation, while the test set only includes words to simulate held-out competition-style data.


<!-- grade_id: Q1_Test, points: 20, grade: true, solution: false -->
```python
# Autograder tests

for subset in dataset:
    assert isinstance(dataset[subset], list), f"Q1: Your {subset} set should be a list. "

for subset, length in zip(dataset.keys(), [7146, 1783, 1516]):
    assert len(dataset[subset]) == length, f"Q1: Your {subset} set is of an incorrect length. "

for subset in dataset:
    for index, sent in enumerate(dataset[subset]):
        assert len(sent) >= 10, f"Q1: Sentence at index {index} of your {subset} set is shorter than required. "
        item_type = tuple if subset in ["train", "dev"] else str
        for item in sent:
            assert isinstance(item, item_type), f"Q1: Your {subset} set should contain lists of {item_type.__name__}s. "

# Some hidden tests below

```

## Question 2: Train an HMM Tagger (10 pts)

Once we have the training, development and testing sets ready, it's straightforward to train an HMM tagger with the help of the `nltk` library. Specifically, check out the [documentation](http://www.nltk.org/api/nltk.tag.html?#nltk.tag.hmm.HiddenMarkovModelTagger) of the `HiddenMarkovModelTagger` class and understand how to use its classmethod `train` to create an HMM tagger. The classmethod `train` also accepts an argument called `test_sequence`, where you can plug in your development data and get a score as an estimate for how well your HMM will perform on unseen data. 

**This function should return a trained `HiddenMarkovModelTagger` object.**

<!-- grade_id: cell-be2a75b95ae4c066, points: 0, grade: false, solution: true -->
```python
from nltk.tag.hmm import HiddenMarkovModelTagger

def train_hmm(train, dev):
    """
    Train an HMM POS tagger on labeled training sentences.
    """
    # the development set is passed as test_sequence so NLTK reports validation accuracy
    hmm_tagger = HiddenMarkovModelTagger.train(
        labeled_sequence = train,
        test_sequence= dev
    )
    
    return hmm_tagger

```

**Commentary:** The HMM learns transition probabilities between POS tags and emission probabilities between POS tags and words from the labeled training sentences. Passing the development set as `test_sequence` lets NLTK immediately estimate validation accuracy without using the hidden-label test set.


```python
# Train an HMM tagger, which takes a while

hmm_tagger = train_hmm(dataset["train"], dataset["dev"])

# The number it reports is the accuracy on the development set
```

<!-- grade_id: Q2_Test, points: 10, grade: true, solution: false -->
```python
# Autograder tests

assert isinstance(hmm_tagger, HiddenMarkovModelTagger), "Q2: Your function should return a HiddenMarkovModelTagger. "
assert hmm_tagger._states, "Q2: Your HMM's states are empty. The tagger wasn't trained properly. "

# Some hidden tests

```

Now that we have trained an HMM tagger, let's now apply it to your testing set to see how it does. For example, the line of code below tags the first sentence in the testing set. Feel free to adapt the code to tag more sentences. 

```python
hmm_tagger.tag(dataset["test"][0])
```

```python
# Output

[('SOCCER', 'NN'),
 ('-', ':'),
 ('JAPAN', 'NNP'),
 ('GET', '.'),
 ('LUCKY', '"'),
 ('WIN', 'NNP'),
 (',', ','),
 ('CHINA', 'NNP'),
 ('IN', 'IN'),
 ('SURPRISE', 'NNP'),
 ('DEFEAT', 'NNP'),
 ('.', '.')]
```